The resource elasticity of control

Author response:

The following is the authors’ response to the previous reviews

Reviewer #1 (Public review):

This research takes a novel theoretical and methodological approach to understanding how people estimate the level of control they have over their environment and how they adjust their actions accordingly. The task is innovative and both it and the findings are well-described (with excellent visuals). They also offer thorough validation for the particular model they develop. The research has the potential to theoretically inform understanding of control across domains, which is a topic of great importance.

We thank the Reviewer for their favorable appraisal and valuable suggestions, which have helped clarify and strengthen the study’s conclusion.

In its revised form, the manuscript addresses most of my previous concerns. The main remaining weakness pertains to the analyses aimed at addressing my suggesting of Bayesian updating as an alternative to the model proposed by the authors. My suggestion was to assume that people perform a form of function approximation to relate resource expenditure to success probability. The authors performed a version of this where people were weighing evidence for a few canonical functions (flat, step, linear), and found that this model underperformed theirs. However, this Bayesian model is quite constrained in its ability to estimate the function relating resources. A more robust test would be to assume a more flexible form of updating that is able to capture a wide range of distributions (e.g., using basis functions, gaussian processes, or nonparametric estimators); see, e.g., work by Griffiths on human function learning). The benefit of testing this type of model is that it would make contact with a known form of inference that individuals engage in across various settings and therefore could offer a more parsimonious and generalizable account of function learning, whereby learning of resource elasticity is a special case. I defer to the authors as to whether they'd like to pursue this direction, but if not I think it's still important that they acknowledge that they are unable to rule out a more general process like this as an alternative to their model. This pertains also to inferences about individual differences, which currently hinge on their preferred model being the most parsimonious.

We thank the Reviewer for this thoughtful suggestion. We acknowledge that more flexible function learning approaches could provide a stronger test in favor of a more general account. Our Bayesian model implemented a basis function approach where the weights of three archetypal functions (flat, step, linear) are learned from experience Testing models with more flexible basis functions would likely require a task with more than three levels of resource investment (1, 2, or 3 tickets). This would make an interesting direction for future work expanding on our current findings. We now incorporate this suggestion in more detail in our updated manuscript (335-341):

Second, future models could enable generalization to levels of resource investment not previously experienced. For example, controllability and its elasticity could be jointly estimated via function approximation that considers control as a function of invested resources. Although our implementation of this model did not fit participants’ choices well (see Methods), other modeling assumptions drawn from human function learning [30] or experimental designs with continuous action spaces may offer a better test of this idea.

Reviewer #2 (Public review):

This research investigates how people might value different factors that contribute to controllability in a creative and thorough way. The authors use computational modeling to try to dissociate "elasticity" from "overall controllability," and find some differential associations with psychopathology. This was a convincing justification for using modeling above and beyond behavioral output and yielded interesting results. Notably, the authors conclude that these findings suggest that biased elasticity could distort agency beliefs via maladaptive resource allocation. Overall, this paper reveals important findings about how people consider components of controllability. The authors have gone to great lengths to revise the manuscript to clarify their definitions of "elastic" and "inelastic" and bolster evidence for their computational model, resulting in an overall strong manuscript that is valuable for elucidating controllability dynamics and preferences.

We thank the Reviewer for their constructive feedback throughout the review process, which has substantially strengthened our manuscript and clarified our theoretical framework.

One minor weakness is that the justification for the analysis technique for the relationships between the model parameters and the psychopathology measures remains lacking given the fact that simple correlational analyses did not reveal any significant associations.

We note that the existence of bivariate relationships is not a prerequisite for the existence of multivariate relationships. Conditioning the latter on the former, therefore, would risk missing out on important relationships existing in the data. Ultimately, correlations between pairs of variables do not offer a sensitive test for the general hypothesis that there is a relationship between two sets of variables. As an illustration, consider that elasticity bias correlated in our data (r = .17, p<.001) with the difference between SOA (sense of agency) and SDS (self-rating depression). Notably, SOA and SDS were positively correlated (r = .47, p<.001), and neither of them was correlated with elasticity bias (SOA: r=.04 p=.43, SDS: r=-.06, p=.16). It was a dimension that ran between them that mapped onto elasticity bias. This specific finding is incidental and uncorrected for multiple comparisons, hence we do not report it in the manuscript, but it illustrates the kinds of relationships that cannot be accounted for by looking at bivariate relationships alone.

Reviewer #3 (Public review):

A bias in how people infer the amount of control they have over their environment is widely believed to be a key component of several mental illnesses including depression, anxiety, and addiction. Accordingly, this bias has been a major focus in computational models of those disorders. However, all of these models treat control as a unidimensional property, roughly, how strongly outcomes depend on action. This paper proposes---correctly, I think---that the intuitive notion of "control" captures multiple dimensions in the relationship between action and outcome.

In particular, the authors identify one key dimension: the degree to which outcome depends on how much *effort* we exert, calling this dimension the "elasticity of control". They additionally argue that this dimension (rather than the more holistic notion of controllability) may be specifically impaired in certain types of psychopathology. This idea has the potential to change how we think about several major mental disorders in a substantial way and can additionally help us better understand how healthy people navigate challenging decision-making problems. More concisely, it is a very good idea.

We thank the Reviewer for their thoughtful engagement with our manuscript. We appreciate their recognition of elasticity as a key dimension of control that has the potential to advance our understanding of psychopathology and healthy decision-making.

Starting with theory, the authors do not provide a strong formal characterization of the proposed notion of elasticity. There are existing, highly general models of controllability (e.g., Huys & Dayan, 2009; Ligneul, 2021) and the elasticity idea could naturally be embedded within one of these frameworks. The authors gesture at this in the introduction; however, this formalization is not reflected in the implemented model, which is highly task-specific.

Our formal definition of elasticity, detailed in Supplementary Note 1, naturally extends the reward-based and information-theoretic definitions of controllability by Huys & Dayan (2009) and Ligneul (2021). We now further clarify how the model implements this formalized definition (lines 156-159).

Conversely, in the ‘elastic controllability model’, the beta distributions represent a belief about the maximum achievable level of control (𝑎Control, 𝑏Control) coupled with two elasticity estimates that specify the degree to which successful boarding requires purchasing at least one (𝑎elastic≥1, 𝑏elastic≥1) or specifically two (𝑎elastic2, 𝑏elastic2) extra tickets. As such, these elasticity estimates quantify how resource investment affects control. The higher they are, the more controllability estimates can be made more precise by knowing how much resources the agent is willing and able to invest (Supplementary Note 1).

Moreover, the authors present elasticity as if it is somehow "outside of" the more general notion of controllability. However, effort and investment are just specific dimensions of action; and resources like money, strength, and skill (the "highly trained birke") are just specific dimensions of state. Accordingly, the notion of elasticity is necessarily implicitly captured by the standard model. Personally, I am compelled by the idea that effort and resource (and therefore elasticity) are particularly important dimensions, ones that people are uniquely tuned to. However, by framing elasticity as a property that is different in kind from controllability (rather than just a dimension of controllability), the authors only make it more difficult to integrate this exciting idea into generalizable models.

We respectfully disagree that we present elasticity as outside of, or different in kind from, controllability. Throughout the manuscript, we explicitly describe elasticity as a dimension of controllability (e.g., lines 70-72, along many other examples). This is also expressed in our formal definition of elasticity (Supplementary Note 1).

The argument that vehicle/destination choice is not trivial because people occasionally didn't choose the instructed location is not compelling to me-if anything, the exclusion rate is unusually low for online studies. The finding that people learn more from non-random outcomes is helpful, but this could easily be cast as standard model-based learning very much like what one measures with the Daw two-step task (nothing specific to control here). Their final argument is the strongest, that to explain behavior the model must assume "a priori that increased effort could enhance control." However, more literally, the necessary assumption is that each attempt increases the probability of success-e.g. you're more likely to get a heads in two flips than one. I suppose you can call that "elasticity inference", but I would call it basic probabilistic reasoning.

We appreciate the Reviewer’s concerns but feel that some of the more subjective comments might not benefit from further discussion. We only note that controllability and its elasticity are features of environmental structure, so in principle any controllability-related inference is a form of model-based learning. The interesting question is whether people account in their world model for that particular feature of the environment.

The authors try to retreat, saying "our research question was whether people can distinguish between elastic and inelastic controllability." I struggle to reconcile this with the claim in the abstract "These findings establish the elasticity of control as a distinct cognitive construct guiding adaptive behavior". That claim is the interesting one, and the one I am evaluating the evidence in light of.

In real-world contexts, it is often trivial that sometimes further investment enhances control and sometimes it does not. For example, students know that if they prepare more extensively for their exams they will likely be able to achieve better grades, but they also know that there is uncertainty in this regard – their grades could improve significantly, modestly, or in some cases, they might not improve at all, depending on the type of exams their study program administers and the knowledge or skills being tested. Our research question was whether in such contexts people learn from experience the degree to which controllability is elastic to invested resources and adapt their resource investment accordingly. Our findings show that they do.

The authors argue for CCA by appeal to the need to "account for the substantial variance that is typically shared among different forms of psychopathology". I agree. A simple correlation would indeed be fairly weak evidence. Strong evidence would show a significant correlation after *controlling for* other factors (e.g. a regression predicting elasticity bias from all subscales simultaneously). CCA effectively does the opposite, asking whether-with the help of all the parameters and all the surveys-one can find any correlation between the two sets of variables. The results are certainly suggestive, but they provide very little statistical evidence that the elasticity parameter is meaningfully related to any particular dimension of psychopathology.

We agree with the Reviewer on the relationship between elasticity and any particular dimension of psychopathology. The CCA asks a different question, namely, whether there is a relationship between psychopathology traits and task parameters, and whether elasticity bias specifically contributes to this relationship.

I am very concerned to see that the authors removed the discussion of this limitation in response to my first review. I quote the original explanation here:

In interpreting the present findings, it needs to be noted that we designed our task to be especially sensitive to overestimation of elasticity. We did so by giving participants free 3 tickets at their initial visits to each planet, which meant that upon success with 3 tickets, people who overestimate elasticity were more likely to continue purchasing extra tickets unnecessarily. Following the same logic, had we first had participants experience 1 ticket trips, this could have increased the sensitivity of our task to underestimation of elasticity in elastic environments. Such underestimation could potentially relate to a distinct psychopathological profile that more heavily loads on depressive symptoms. Thus, by altering the initial exposure, future studies could disambiguate the dissociable contributions of overestimating versus underestimating elasticity to different forms of psychopathology.

The logic of this paragraph makes perfect sense to me. If you assume low elasticity, you will infer that you could catch the train with just one ticket. However, when elasticity is in fact high, you would find that you don't catch the train, leading you to quickly infer high elasticity eliminating the bias. In contrast, if you assume high elasticity, you will continue purchasing three tickets and will never have the opportunity to learn that you could be purchasing only one-the bias remains.

The authors attempt to argue that this isn't happening using parameter recovery. However, they only report the *correlation* in the parameter, whereas the critical measure is the *bias*. Furthermore, in parameter recovery, the data-generating and data-fitting models are identical-this will yield the best possible recovery results. Although finding no bias in this setting would support the claims, it cannot outweigh the logical argument for the bias that they originally laid out. Finally, parameter recovery should be performed across the full range of plausible parameter values; using fitted parameters (a detail I could only determine by reading the code) yields biased results because the fitted parameters are themselves subject to the bias (if present). That is, if true low elasticity is inferred as high elasticity, then you will not have any examples of low elasticity in the fitted parameters and will not detect the inability to recover them.

The logic the Reviewer describes breaks down when one considers the dynamics of participants’ resource investment choices. A low elasticity bias in a participant’s prior belief would make them persist for longer in purchasing a single ticket despite failure, as compared to a person without such a bias. Indeed, the ability of the experimental design to demonstrate low elasticity biases is evidenced by the fact that the majority of participants were fitted with a low elasticity bias (μ = .16 ± .14, where .5 is unbiased).

Originally, the Reviewer was concerned that elasticity bias was being confounded with a general deficit in learning. The weak inter-parameter correlations in the parameter recovery test resolved this concern, especially given that, as we now noted, the simulated parameter space encompassed both low and high elasticity biases (range=[.02,.76]). Furthermore, regarding the Reviewer's concern about bias in the parameter recovery, we found no such significant bias with respect to the elasticity bias parameter (Δ(Simulated, Recovered)= -.03, p=.25), showing that our experiment could accurately identify low and high elasticity biases.

The statistical structure of the task is inconsistent with the framing. In the framing, participants can make either one or two second boarding attempts (jumps) by purchasing extra tickets. The additional attempt(s) will thus succeed with probability p for one ticket and 2p – p^2 for two tickets; the p^2 captures the fact that you only take the second attempt if you fail on the first. A consequence of this is buying more tickets has diminishing returns. In contrast, in the task, participants always jumped twice after purchasing two tickets, and the probability of success with two tickets was exactly double that with one ticket. Thus, if participants are applying an intuitive causal model to the task, they will appear to "underestimate" the elasticity of control. I don't think this seriously jeopardizes the key results, but any follow-up work should ensure that the task's structure is consistent with the intuitive causal model.

We thank the Reviewer for this comment, and agree the participants may have employed the intuitive understanding the Reviewer describes. This is consistent with our model comparison results, which showed that participants did not assume that control increases linearly with resource investment (lines 677-692). Consequently, this is also not assumed by our model, except perhaps by how the prior is implemented (a property that was supported by model comparison). In the text, we acknowledge that this aspect of the model and participants’ behavior deviates from the true task's structure, and it would be worthwhile to address this deviation in future studies.

That said, there is no reason that this will make participants appear to be generally underestimating elasticity. Following exposure to outcomes for one and three tickets, any nonlinear understanding of probabilities would only affect the controllability estimate for two tickets. This would have contrasting effects on the elasticity estimated to the second and third tickets, but on average, it would not change the overall elasticity estimated. On the other hand, such a participant is only exposed to outcomes for two and three tickets, they would come to judge the difference between the first and second tickets too highly, thereby overestimating elasticity.

The model is heuristically defined and does not reflect Bayesian updating. For example, it overestimates maximum control by not using losses with less than 3 tickets (intuitively, the inference here depends on what your beliefs about elasticity). Including forced three-ticket trials at the beginning of each round makes this less of an issue; but if you want to remove those trials, you might need to adjust the model. The need to introduce the modified model with kappa is likely another symptom of the heuristic nature of the model updating equations.

Note that we have tested a fully Bayesian model (lines 676-691), but found that this model fitted participants’ choices worse.

You're right; saying these analyses provides "no information" was unfair. I agree that this is a useful way to link model parameters with behavior, and they should remain in the paper. However, my key objection still holds: these analyses do not tell us anything about how *people's* prior assumptions influence behavior. Instead, they tell us about how *fitted model parameters* depend on observed behavior. You can easily avoid this misreading by adding a small parenthetical, e.g.

Thus, a prior assumption that control is likely available **(operationalized by \gamma_controllability)** was reflected in a futile investment of resources in uncontrollable environments.

We thank the Reviewer for the suggestion and have added this parenthetical (lines 219, 225).

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The ability to determine how much the environment can be controlled through our actions has long been viewed as fundamental to adaptive behavior. While traditional accounts treat controllability as a fixed property of the environment, we argue that real-world controllability often depends on the effort, time and money we are able and willing to invest. In such cases, controllability can be said to be elastic to invested resources. Here we propose that inferring this elasticity is essential for efficient resource allocation, and thus, elasticity misestimations result in maladaptive behavior. To test this hypothesis, we developed a novel treasure hunt game where participants encountered environments with varying degrees of controllability and elasticity. Across two pre-registered studies (N=514), we first demonstrate that people infer elasticity and adapt their resource allocation accordingly. We then present a computational model that explains how people make this inference, and identify individual elasticity biases that lead to suboptimal resource allocation. Finally, we show that overestimation of elasticity is associated with elevated psychopathology involving an impaired sense of control. These findings establish the elasticity of control as a distinct cognitive construct guiding adaptive behavior, and a computational marker for control-related maladaptive behavior.

Revised version: The ability to determine how much the environment can be controlled through our actions has long been viewed as fundamental to adaptive behavior. While traditional accounts treat controllability as a fixed property of the environment, we argue that real-world controllability often depends on the effort, time and money we are able and willing to invest. In such cases, controllability can be said to be elastic to invested resources. Here we propose that inferring this elasticity is essential for efficient resource allocation, and thus, elasticity misestimations result in maladaptive behavior. To test this hypothesis, we developed a novel treasure hunt game where participants encountered environments with varying degrees of controllability and elasticity.

Added in the revised version: Across two pre-registered studies (N=514), we first demonstrate that people infer elasticity and adapt their resource allocation accordingly. We then present a computational model that explains how people make this inference, and identify individual elasticity biases that lead to suboptimal resource allocation. Finally, we show that overestimation of elasticity is associated with elevated psychopathology involving an impaired sense of control. These findings establish the elasticity of control as a distinct cognitive construct guiding adaptive behavior, and a computational marker for control-related maladaptive behavior.

Introduction

Taking action requires us to invest resources, such as time, money, and effort, and is thus only sensible if our actions have the potential to change the environment to our advantage. Accurate estimation of this potential – that is, of the environment’s controllability1–4 – enables us to make informed decisions as to whether to engage in action5–8. Misestimations of controllability thus result in maladaptive behavior, and contribute to mental health disorders such as depression, anxiety, and obsessive-compulsive disorder9–13.

Earlier version: The degree of control we possess over our environment, however, may itself depend on the resources we are willing and able to invest. For example, the control a biker has over their commute time depends on the power they are willing and able to invest in pedaling. In this respect, a highly trained biker would typically have more control than a novice. Likewise, the control a diner in a restaurant has over their meal may depend on how much money they have to spend. In such situations, controllability is not fixed but rather elastic to available resources (i.e., in the same sense that supply and demand may be elastic to changing prices14).

Revised version: The degree of control we possess over our environment, however, may itself depend on the resources we are willing and able to invest. For example, the control a biker has over their commute time depends on the power they are willing and able to invest in pedaling. In this respect, a highly trained biker would typically have more control than a novice. Likewise, the control a diner in a restaurant has over their meal may depend on how much money they have to spend. In such situations, controllability is not fixed but rather elastic to available resources (i.e., in the same sense that supply and demand may be elastic to changing prices14). To formalize how elasticity relates to control, we build on an established definition of controllability as the fraction of reward that is controllably achievable15, χ. Uncertainty about this fraction could result from uncertainty about the amount of resources that the agent is able and willing to invest, max C. Elasticity can thus be defined as the amount of information obtained about controllability by knowing the amount of available resources: I(χ; max C).

Removed in the revised version: To formalize how elasticity relates to control, we build on an established definition of controllability as the fraction of reward that is controllably achievable15, χ. Uncertainty about this fraction could result from uncertainty about the amount of resources that the agent is able and willing to invest, max C. Elasticity can thus be defined as the amount of information obtained about controllability by knowing the amount of available resources: I(χ; max C).

While only controllable environments can be elastic, the inverse is not necessarily true – controllability can be high, yet inelastic to invested resources – for example, choosing between bus routes affords equal control over commute time to anyone who can afford the basic fare (Figure 1; Supplementary Note 1). That said, since all actions require some resource investment, no controllable environment is completely inelastic when considering the full spectrum of possible agents, including those with insufficient resources to act (e.g., those unable to purchase a bus fare or pay for a fixed-price meal). Even in this case, however, controllability can be elastic to varying degrees.

Examples of inelastic vs elastic control: Top –

Choosing one of three equally priced public transport routes provides inelastic control over commute time, as control does not increase with additional available resources beyond the basic fare. In contrast, the control of a biker over commute time is elastic to the power they can invest in pedaling. Middle – A diner has either inelastic or elastic control over their lunch depending on whether the restaurant offers a fixed-price lunch deal where all diners have the same options, or a standard menu where more expensive dishes are only available to those able to spend more. Bottom – the probability of winning in blackjack is inelastic to money, as it only depends on whether one hits or stands, whereas when playing roulette, one can increase the probability of winning by investing more coins to cover more possible outcomes. Importantly, most real-world scenarios lie on a spectrum between these illustrative extremes, such that controllability is partly elastic and partly inelastic. In all situations, agents must infer the degree to which controllability is elastic to be able to determine whether the potential gains in control outweigh the costs of investing additional resources (e.g., physical exertion, money spent, time invested).

Earlier version: Knowing the elasticity of control is necessary for deciding whether and how to act. Failing to achieve one’s goal, for example, should only lead one to invest more resources in an elastic environment. Critically, in many real-life domains (e.g., studying for an exam), we must use our experiences to infer the degree to which we can improve our control by investing more resources16–22. Failure to do so is bound to result in maladaptive behavior, and may thus contribute to mental health disorders. Depressed patients, for instance, are less willing to expend their resources to obtain control23,28–29, which could reflect an underestimation of elasticity.

Revised version: Knowing the elasticity of control is necessary for deciding whether and how to act. Failing to achieve one’s goal, for example, should only lead one to invest more resources in an elastic environment. Critically, in many real-life domains (e.g., studying for an exam), we must use our experiences to infer the degree to which we can improve our control by investing more resources16–22. Failure to do so is bound to result in maladaptive behavior, and may thus contribute to mental health disorders. Depressed patients, for instance, are less willing to expend their resources to obtain control23,28–29, which could reflect an underestimation of elasticity. Conversely, the repetitive and unusual amount of effort invested by people with obsessive-compulsive disorder in attempts to exert control25,26 could indicate an overestimation of elasticity, that is, a belief that adequate control can only be achieved through excessive and repeated resource investment26. Ultimately, not having a good grasp of how additional efforts influence the success of one’s actions is bound to undermine one’s sense of agency16–17.

Removed in the revised version: Conversely, the repetitive and unusual amount of effort invested by people with obsessive- compulsive disorder in attempts to exert control25,26 could indicate an overestimation of elasticity, that is, a belief that adequate control can only be achieved through excessive and repeated resource investment26. Ultimately, not having a good grasp of how additional efforts influence the success of one’s actions is bound to undermine one’s sense of agency16–17.

Thus, here we ask whether and how people infer the elasticity of control over their environment. We hypothesize that individual differences in elasticity inference lead to a mismanagement of one’s resources when attempting to control the environment, and are thus associated with psychopathologies involving a distorted sense of control23–26.

Experimental paradigms to date have conflated overall controllability and its elasticity, such that controllability was either low or elastic16–22. The elasticity of control, however, must be dissociated from overall controllability to accurately diagnose mismanagement of resources. A given individual, for instance, may tend to overinvest resources because they overestimate controllability – for example, exercising due to a misguided belief that this can make one taller, when in fact height cannot be controlled. Alternatively, they may do so because they overestimate the elasticity of control – for example, a chess expert practicing unnecessarily hard to win against a novice, when their existing skill level already ensures control over the match’s outcome. Identifying which misestimation drives resource mismanagement requires an experimental design that independently manipulates overall controllability from its elasticity.

Earlier version: Therefore, here we dissociate elastic and inelastic controllability using a novel treasure-hunt game where participants encountered three distinct environments: The first offered high elastic controllability, where a high level of control can only be attained by investing extra resources. The second offered high inelastic controllability, where control is attainable with any level of resource investment. And the third offered low controllability, where control was not attainable despite any level of resource investment. In each environment, participants were allowed to invest various amounts of resources to pursue reward or to altogether forgo pursuing reward.

Revised version: Therefore, here we dissociate elastic and inelastic controllability using a novel treasure-hunt game where participants encountered three distinct environments: The first offered high elastic controllability, where a high level of control can only be attained by investing extra resources. The second offered high inelastic controllability, where control is attainable with any level of resource investment. And the third offered low controllability, where control was not attainable despite any level of resource investment. In each environment, participants were allowed to invest various amounts of resources to pursue reward or to altogether forgo pursuing reward. Examining how participants adapted their resource investment to observed outcomes enabled us to determine whether and how participants inferred each environment’s controllability and its elasticity.

Removed in the revised version: Examining how participants adapted their resource investment to observed outcomes enabled us to determine whether and how participants inferred each environment’s controllability and its elasticity.

Across two pre-registered studies (N = 514), we establish that people do indeed infer the elasticity of control. We present a computational model that explains how people make this inference, and identify individual biases in its implementation. We thus show that elasticity overestimation leads to an overinvestment of resources and is associated with elevated psychopathology involving an impaired sense of control. These findings establish elasticity as a crucial dimension of controllability that guides adaptive behavior, and a computational marker of control-related psychopathology.

Results

Earlier version: We had online participants (first study = 264, replication study = 250, ages =18-45, Mean 33 ± .6) play a novel treasure-hunt game set across multiple planets. On each trip (trial), participants attempted to travel from an initial location (‘desert’ or ‘fountain’) to a treasure (worth 150 coins; 15¢ of monetary compensation), located at either the house or the mountain (Figure 2A).

Revised version: We had online participants (first study = 264, replication study = 250, ages =18-45, Mean 33 ± .6) play a novel treasure-hunt game set across multiple planets. On each trip (trial), participants attempted to travel from an initial location (‘desert’ or ‘fountain’) to a treasure (worth 150 coins; 15¢ of monetary compensation), located at either the house or the mountain (Figure 2A). Participants could exercise control by boarding the train to reach the house, or the plane to reach the mountain (right-side Figure 2B). The present planet’s controllability determined the probability that participants would succeed in boarding the vehicle they selected. In high-controllability planets, participants consistently boarded their preferred vehicle, enabling them to reliably reach their destination. Conversely, in low-controllability planets, participants frequently failed to board and were thus forced to walk to the nearest location. On 20% of trials, the nearest location was where the treasure was located, thus partially decoupling success in exerting control from reward.

Experimental design:

(A) Goal. On each trip to a planet, participants’ goal was to reach a treasure located either at the house or the mountain (B) Transition rules. Participants could exercise control by boarding either the plane or the train to a destination of their choice, whereas missing the ride sent the participant walking to the nearest location (which happened to be the treasure location in only 20% of trials). (C) Trip structure. At the beginning of each trip, participants selected whether to purchase 1, 2, or 3 tickets to attempt to board their vehicle of choice, or walk for free to the nearest location. If the participant purchased at least one ticket, they were allowed to choose between the plane and the train. Then, for each additional ticket purchased, the participant was given an opportunity to increase their chances of boarding the vehicle. Specifically, the chosen vehicle appeared moving from left to right across the screen, and the participant attempted to board it by pressing the spacebar when it reached the center of the screen. At the end of the trip, participants were shown where they arrived at, allowing them to infer whether they successfully boarded the vehicle. (D) The space of possible planets. Planets varied in inelastic controllability – the probability that even one ticket would lead to successfully boarding the vehicle – and in elastic controllability – the degree to which two extra tickets would improve the probability of successfully boarding the vehicle (one extra ticket provided half the benefit). The color corresponds to the optimal number of tickets (0, 1 or 3) in each planet. The darker the color the higher the advantage in expected value gained by purchasing the optimal number of tickets relative to the second-best option. Each participant made 30 consecutive trips in one planet from the green area, one planet from the blue area, one planet from the red area, and one planet with identical characteristic across all participants (black circle).

Removed in the revised version: Participants could exercise control by boarding the train to reach the house, or the plane to reach the mountain (right-side Figure 2B). The present planet’s controllability determined the probability that participants would succeed in boarding the vehicle they selected. In high- controllability planets, participants consistently boarded their preferred vehicle, enabling them to reliably reach their destination. Conversely, in low-controllability planets, participants frequently failed to board and were thus forced to walk to the nearest location. On 20% of trials, the nearest location was where the treasure was located, thus partially decoupling success in exerting control from reward.

To investigate how people learn the elasticity of control, we allowed participants to invest different amounts of resources in attempting to board their preferred vehicle. Participants could purchase one (40 coins), two (60 coins), or three tickets (80 coins) or otherwise walk for free to the nearest location. Participants were informed that a single ticket allowed them to board only if the vehicle stopped at the station, while additional tickets provided extra chances to board even after the vehicle had left the platform. For each additional ticket, the chosen vehicle appeared moving from left to right across the screen, and participants could attempt to board it by pressing the spacebar when it reached the center of the screen. Thus, each additional ticket could increase the chance of boarding but also required a greater investment of resources—decreasing earnings, extending the trial duration, and demanding attentional effort to precisely time a button press when attempting to board.

Participants received no feedback about individual boarding attempts but were shown their initial and final locations at the end of each trip, enabling them to infer whether they had successfully boarded their vehicle. For simplicity, we defined inelastic controllability as the probability that even one ticket would lead to successfully boarding the vehicle, and elastic controllability as the degree to which two extra tickets would increase that probability.

To dissociate elastic from inelastic controllability, each participant visited three randomly sampled planets. In the first planet, controllability was high and inelastic, meaning that no matter how many tickets you purchased, you had a good chance of making your ride (sampled from green area in Figure 2D). In the second planet, controllability was high and elastic, meaning that you needed to purchase two extra tickets to have a good chance of making your ride (Figure 2D, blue). And in the third planet, controllability was low, meaning that no matter how many tickets you purchased, you did not have a good chance of making your ride (Figure 2D, red).

Importantly, the order of these planets was randomized, and participants did not know in advance which planet they were visiting. Thus, they could only infer a planet’s controllability and elasticity from the outcomes of their actions. Moreover, to increase sensitivity to individual biases, we had all participants visit a fourth planet where purchasing any number of tickets (0, 1, 2, 3) was equally optimal (Figure 2D, black circle). We also homogenized participants’ initial learning experiences by giving them three free tickets for the first five visits to every planet. As we shall see below, this latter measure also helped dissociate between different models of how participants solved the task.

People adapt to the elasticity of control

To determine whether participants inferred the elasticity of control, we first examined how they adapted their resource investment across planets. We found that participants were more likely to purchase tickets (‘opt in’) in more controllable planets, whether controllability was elastic (Initial and Replication studies, p<.001; mixed logistic regression) or inelastic (Initial and Replication studies, p <.001; Figure 3A&C). Conversely, participants purchased more extra tickets only in planets with high elastic controllability (Initial p =.01, Replication p=.004; mixed probit regression; Figure 3B&C). Although substantial individual differences were observed in this regard (Figure 3D), these results show that most participants successfully adapted to both the controllability of the environments they visited and its elasticity.

Participants adapted to the elasticity of control:

Results from initial (N = 264) and replication (N = 250) studies. (A) Opt-in percentage across all planets. Participants opted-in more frequently on controllable planets (outlined in blue and green) than on uncontrollable planets (outlined in red) (B) Extra tickets purchased across all planets. Participants purchased more extra tickets in planets with high elastic controllability (outlined in blue). An average of 11 to 12 participants visited each planet in each of the studies. (C) Effect of elastic and inelastic controllability on opting in and the purchasing of extra-tickets. Bars show estimated fixed effects and 95% CIs from mixed logistic (opt-in) and probit (extra tickets) regressions. (D) Individual differences. Distribution of participants by opt-in rate (top) and average number of extra tickets purchased (bottom) across both studies, shown separately for planets with low controllability (red), high inelastic controllability (green), and high elastic controllability (blue).

People infer the elasticity of control

That participants adapted to each environment does not necessarily mean they inferred the elasticity of control (nor that they inferred control at all, but see Methods for how this was verified). Successful adaptation to the elasticity of control may also be achieved merely by treating each possible level of resource investment as a separate course of action. One could thus learn the degree of control afforded by each course of action (i.e., purchasing 1, 2, or 3 tickets), and then choose the one that affords the highest expected value considering its cost. This simple strategy could achieve optimal performance given unlimited experience. However, it would not be as efficient as a strategy that employs elasticity inference, because it does not make use of the dependencies that exist between different levels of resource investment in the degrees of control they afford. Thus, upon failing to board a vehicle despite purchasing 3 tickets, the simpler strategy would need to purchase 1 or 2 tickets to know these options would fail as well, whereas an agent equipped with the concept of elasticity would conclude that controllability is low, and it is thus best to opt out. Conversely, upon succeeding to board with 3 tickets, the simpler strategy would tend to favor purchasing 3 tickets, over 1 or 2 tickets, even if the extra tickets confer no benefit, whereas an agent that infers elasticity would not favor 3 over 1 or 2 tickets because it received no evidence that controllability is elastic (i.e., that the extra tickets help).

Earlier version: To determine which of these models best describes the way participants solved the task, we formalized the simpler strategy as the ‘controllability model’, and an elasticity-inferring strategy as the ‘elastic controllability model’ (pre-registered at: https://aspredicted.org/CHW_12H). Both models use beta distributions, characterized by parameters a and b, to represent the probability of boarding when purchasing different numbers of tickets, but only the elastic controllability model learns latent elasticity estimates. Specifically, in the ‘controllability model’, three beta distributions represent the expected level of control from purchasing 1, 2, or 3 tickets. Thus, the parameters of these distributions (an, bn, n ∈ {1, 2, 3}) accumulate the number of times each number of tickets led (an) or did not lead (bn) to successful boarding. Conversely, in the ‘elastic controllability model’, the beta distributions represent a belief about the maximum achievable level of control (acontrol, bcontrol) coupled with two elasticity estimates that specify the degree to which successful boarding requires purchasing at least one (aelastic≥1,, belastic≥1,) or specifically two (aelastic2, belastic2) extra tickets. Figure 4 shows schematically how the two models differ in updating their estimates in response to different types of outcomes. Both models similarly use these estimates to compute an expected value for purchasing 0, 1, 2, or 3 tickets, and accordingly choose between these options (see Methods for model equations and Supplementary Table 3 for all parameter fits, priors, and their interpretation).

Revised version: To determine which of these models best describes the way participants solved the task, we formalized the simpler strategy as the ‘controllability model’, and an elasticity-inferring strategy as the ‘elastic controllability model’ (pre-registered at: https://aspredicted.org/CHW_12H). Both models use beta distributions, characterized by parameters a and b, to represent the probability of boarding when purchasing different numbers of tickets, but only the elastic controllability model learns latent elasticity estimates. Specifically, in the ‘controllability model’, three beta distributions represent the expected level of control from purchasing 1, 2, or 3 tickets. Thus, the parameters of these distributions (𝑎n, 𝑏n, n ∈ {1, 2, 3}) accumulate the number of times each number of tickets led (𝑎n) or did not lead (𝑏n) to successful boarding. Conversely, in the ‘elastic controllability model’, the beta distributions represent a belief about the maximum achievable level of control (𝑎control, 𝑏control) coupled with two elasticity estimates that specify the degree to which successful boarding requires purchasing at least one (𝑎elastic≥1, 𝑏elastic≥1) or specifically two (𝑎elastic2, 𝑏elastic2) extra tickets. As such, these elasticity estimates quantify how resource investment affects control. The higher they are, the more controllability estimates can be made more precise by knowing how much resources the agent is willing and able to invest (Supplementary Note 1). Figure 4 shows schematically how the two models differ in updating their estimates in response to different types of outcomes. Both models similarly use these estimates to compute an expected value for purchasing 0, 1, 2, or 3 tickets, and accordingly choose between these options (see Methods for model equations and Supplementary Table 3 for all parameter fits, priors, and their interpretation).

Computational models:

Earlier version: Illustration of the learning rules of the controllability and elastic controllability models. The width of the colored region represents the estimated control, shown as a percentage of absolute control (gray area) along with the update rules based on each outcome. (A) Controllability model: This model treats the purchase of 1, 2, and 3 tickets as distinct actions. It accumulates evidence for the effectiveness (a) or lack of effectiveness (b) of each action based solely on whether or not the action led to successful boarding, illustrated by changes in the width of the shaded region corresponding to each ticket amount. (B) Elastic Controllability model: Represents beliefs about maximum controllability (black outline) and the degree to which at least one or two extra tickets are necessary to obtain it. These beliefs are used to calculate the expected control when purchasing 1 ticket (inelastic controllability) and the additional control afforded by 2 and 3 tickets (elastic controllability), demonstrated by the bars. The overlapping arrangement of the bars illustrates that, only in the elastic controllability model, maximum controllability is the sum of inelastic and elastic control. The model updates its beliefs as follows: Top – Successfully boarding with one ticket provides evidence of high maximum controllability (acontrol + 1; expanded shaded region) and reduces the perceived need to purchase extra tickets (belastic≥1 + 1; light green expanded). A failure to board does not change estimated maximum controllability, but rather suggests that 1 ticket might not suffice to obtain control (aelastic≥1 + 1; light green diminished). As a result, the model’s estimate of control afforded by 1 ticket, as well as on average across ticket options, is reduced. Middle – Successfully boarding with 2 tickets provides evidence of high maximum controllability (acontrol + 1; expanded shaded region), and reduces the perceived need to purchase two extra tickets (belastic2 + 1; light & dark green expanded). Here too, a failure to board does not change estimated maximum controllability, but rather suggests that 2 tickets might not be sufficient, thus reducing the estimated control for 2 tickets and on average (aelastic2 + 1; light & dark green diminished). Bottom – Successfully boarding with 3 tickets provides evidence of high maximum controllability (acontrol + 1; expanded shaded region), but offers no insight about whether fewer tickets would suffice to obtain it. Failure to board with 3 tickets provides evidence of low controllability (bcontrol + 1; diminished shaded region).

Revised version: Illustration of the learning rules of the controllability and elastic controllability models. The width of the colored region represents the estimated control, shown as a percentage of absolute control (gray area) along with the update rules based on each outcome. (A) Controllability model: This model treats the purchase of 1, 2, and 3 tickets as distinct actions. It accumulates evidence for the effectiveness (𝑎) or lack of effectiveness (𝑏) of each action based solely on whether or not the action led to successful boarding, illustrated by changes in the width of the shaded region corresponding to each ticket amount. (B) Elastic Controllability model: Represents beliefs about maximum controllability (black outline) and the degree to which at least one or two extra tickets are necessary to obtain it. These beliefs are used to calculate the expected control when purchasing 1 ticket (inelastic controllability) and the additional control afforded by 2 and 3 tickets (elastic controllability), demonstrated by the bars. The overlapping arrangement of the bars illustrates that, only in the elastic controllability model, maximum controllability is the sum of inelastic and elastic control. The model updates its beliefs as follows: Top – Successfully boarding with one ticket provides evidence of high maximum controllability (𝑎control + 1; expanded shaded region) and reduces the perceived need to purchase extra tickets (𝑏elastic≥1 + 1; light green expanded). A failure to board does not change estimated maximum controllability, but rather suggests that 1 ticket might not suffice to obtain control (𝑎elastic≥1 + 1; light green diminished). As a result, the model’s estimate of control afforded by 1 ticket, as well as on average across ticket options, is reduced. Middle – Successfully boarding with 2 tickets provides evidence of high maximum controllability (𝑎control + 1; expanded shaded region), and reduces the perceived need to purchase two extra tickets (𝑏elastic2 + 1; light & dark green expanded). Here too, a failure to board does not change estimated maximum controllability, but rather suggests that 2 tickets might not be sufficient, thus reducing the estimated control for 2 tickets and on average (𝑎elastic2 + 1; light & dark green diminished). Bottom – Successfully boarding with 3 tickets provides evidence of high maximum controllability (𝑎control + 1; expanded shaded region), but offers no insight about whether fewer tickets would suffice to obtain it. Failure to board with 3 tickets provides evidence of low controllability (𝑏control + 1; diminished shaded region).

Examining the predictions of both models, each simulated using the parameter setting that best fit participants’ choices, showed that the elastic-controllability model was better than the controllability model at opting-out in low controllability planets (elastic-controllability model % opt in = 49% ± .01, controllability model % opt in = 63% ± .01; Figure 5C, top panel). Most importantly, this higher level of performance closely matched participants’ behavior (% opt in = 47%, ±.02). Furthermore, the elastic controllability model was also better at not purchasing extra tickets in low-controllability (elastic controllability extra tickets = .49±.01 controllability extra tickets = .65±.01) and inelastic-controllability (elastic controllability extra tickets = .58±.01 controllability extra tickets = .66±.01) planets, and this too was more closely aligned with participants’ behavior (low-controllability extra tickets = 51±.02, inelastic-controllability extra tickets = 53±.02; Figure 5C, bottom panel).

Model simulations:

(A) Opt-in. Proportion of opt-in by the controllability and elastic controllability models across all planets. Each region is outlined by the color corresponding to the optimal ticket strategy (Blue: 3 tickets, Green: 1 ticket, Red: opt-out). (B) Extra ticket purchases. Extra tickets purchased by the controllability and elastic controllability models across all planets. Here too, each region is outlined by the color corresponding to the optimal ticket strategy. (C) Comparison of opt-in rates and extra ticket purchases made by models and participants. Top - Proportion of opting-in as a function of overall controllability. Bottom - The number of extra tickets purchased for each level of inelastic controllability when elasticity did not warrant purchasing extra tickets (i.e. below 30%). (D) Resource investment following free 3-ticket trips. The distribution of resource investment choices made immediately following the free 3-ticket trips, as a function of the outcomes of those trips. Results are shown for the two studies (Initial, Replication), the controllability model, and the original and modified variant of the elastic controllability models. (E) Model comparison. Consistency of each model with participant behavior is shown as log Bayes factor compared to a null model wherein participants have distinct preferences among different number of tickets to purchase but no latent controllability estimates.

Beyond these general aspects of task behavior, having participants start on every planet with free 3-ticket trips also enabled us to examine their response to succeeding or failing with 3 tickets, prior to any other experiences on the planet (Figure 5D). As noted above, failure with 3 tickets makes the controllability model favor 1 and 2 tickets over 3 tickets, whereas success with 3 tickets makes it favor 3 tickets over 1 and 2 tickets. Critically, neither the participants nor the simulation of the elastic controllability model showed these effects. Upon failing with 3 tickets, both participants and the elastic model primarily favored to opt out. Conversely, if they succeeded with 3 tickets, they did not become less likely to purchase 1 or 2 tickets.

In fact, failure with 3 tickets even made participants favor 3, over 1 and 2, tickets. This favoring of 3 tickets continued until participants accumulated sufficient evidence about their limited control to opt out (Supplementary Figure 2). Presumably, the initial failures with 3 tickets resulted in an increased uncertainty about whether it is at all possible to control one’s destination. Consequently, participants who nevertheless opted in invested maximum resources to resolve this uncertainty before exploring whether control is elastic. While this behavior directly contradicts the controllability model, it also goes beyond the behavior exhibited by the elastic controllability model. To capture this aspect of participants’ behavior, we modified the elastic controllability model (elastic controllability*) such that it would increase its elasticity estimates when controllability is reduced. Model comparison consistently supported this modification (Initial study: log Bayes Factor = 143 Replication study: log Bayes factor = 155).

More importantly, both the original and modified elastic controllability models explained participants’ behavior better than the controllability model (Initial: log Bayes Factor ≥ 270, Replication: log Bayes Factor ≥ 148; Figure 5E) as well as other theoretically motivated alternatives (see Methods). These results establish that people infer the elasticity of control, and thereby adapt their resource allocation to the environment’s controllability and its elasticity.

Individual elasticity biases explain misallocation of resources

Earlier version: To determine whether the elastic controllability* model could identify individual biases in elasticity inference, we fitted the model to participants’ choices and thus derived each participant’s beliefs about controllability (γcontrollability) and elasticity (γelasticity) prior to visiting any planet. γcontrollability is used to initialize the model’s belief about the maximum available control (acontrol, bcontrol), with higher values reflecting stronger prior beliefs that planets are controllable. γelasticity is used to initialize the model’s elasticity estimates (aelastic, belastic), with higher values reflecting stronger prior beliefs that control is elastic (see Methods). The stronger the belief the greater its influence on choices to opt in (γcontrollability) and purchase extra tickets (γelasticity) regardless of observed outcomes. In addition to these two model parameters, we examined four additional parameters that directly influence opting-in and the purchase of extra tickets. These include an inverse temperature parameter (β), which specifies the stochasticity with which expected values affect resource investment choices, and baseline- preference parameters for purchasing 1, 2, or 3 tickets (α1, α2, α3). Across both samples, shared variance among parameters never exceeded 14% (Supplementary Figure 1), and was particularly low for the controllability and elasticity priors (Initial study: 5%, Replication study: 3%). Thus, a tendency to assume the environment is controllable did not imply a tendency to assume that control is elastic, establishing inferences concerning maximum controllability and it elasticity as distinct facets of individual differences.

Revised version: To determine whether the elastic controllability* model could identify individual biases in elasticity inference, we fitted the model to participants’ choices and thus derived each participant’s beliefs about controllability (𝛾controllability) and elasticity (𝛾elasticity) prior to visiting any planet. 𝛾controllability is used to initialize the model’s belief about the maximum available control (𝑎control, 𝑏control), with higher values reflecting stronger prior beliefs that planets are controllable. 𝛾elasticity is used to initialize the model’s elasticity estimates (𝑎elastic, 𝑏elastic), with higher values reflecting stronger prior beliefs that control is elastic (see Methods). The stronger the belief the greater its influence on choices to opt in (𝛾controllability) and purchase extra tickets (𝛾elasticity) regardless of observed outcomes. In addition to these two model parameters, we examined four additional parameters that directly influence opting-in and the purchase of extra tickets. These include an inverse temperature parameter (β), which specifies the stochasticity with which expected values affect resource investment choices, and baseline-preference parameters for purchasing 1, 2, or 3 tickets (α1, α2, α3). Across both samples, shared variance among parameters never exceeded 14% (Supplementary Figure 1), and was particularly low for the controllability and elasticity priors (Initial study: 5%, Replication study: 3%). Thus, a tendency to assume the environment is controllable did not imply a tendency to assume that control is elastic, establishing inferences concerning maximum controllability and it elasticity as distinct facets of individual differences.

Earlier version: To clarify how fitted model parameters related to observable behavior, we regressed participants’ opt-in rates and extra ticket purchases on the parameters (Figure 6A). We found that a higher controllability bias (higher γcontrollability) was associated with more opting in (Initial study p<.001, Replication study p<.02; logistic regression) and more extra tickets purchased (Initial p<.001, Replication p<.001; probit regression) on low-controllability, as compared to high- controllability, planets. Thus, a prior assumption that control is likely available was reflected in a futile investment of resources in uncontrollable environments.

Revised version: To clarify how fitted model parameters related to observable behavior, we regressed participants’ opt-in rates and extra ticket purchases on the parameters (Figure 6A). We found that a higher controllability bias (higher 𝛾controllability) was associated with more opting in (Initial study p<.001, Replication study p<.02; logistic regression) and more extra tickets purchased (Initial p<.001, Replication p<.001; probit regression) on low-controllability, as compared to high-controllability, planets. Thus, a prior assumption that control is likely available (operationalized by 𝛾controllability) was reflected in a futile investment of resources in uncontrollable environments.

Added in the revised version: (A) Model parameters and task behavior.

Earlier version: (A) Model parameters and task behavior. Model parameters fitted to each participant’s task behavior were used as predictors of opt-in (left; multiple logistic regressions) and extra ticket purchases (right; multiple probit regressions) separately for low controllability planets where opting out is optimal (red), high elastic controllability planets where purchasing 3 tickets is optimal (blue), and high inelastic controllability planets where purchasing 1 ticket is optimal (green). In a separate analysis, participants’ composite psychopathology scores (see panel C) were similarly used to predict opt-in (left; logistic regression) and extra ticket purchases (right; probit regression). Bars represent regression coefficients and 95% CIs. (B) Model parameters and psychopathology. Loadings from a Canonical Correlation Analysis (CCA) between model parameters and self-report psychopathology measures. Self- report measures were scored such that higher values reflect higher levels of psychopathology. Loadings are shown for the Initial (dark blue), Replication (light blue), and combined (gray bars) datasets. The magnitude of the loading reflects the degree to which each measure contributes to the correlation between parameters and self-reports. (C) Model parameters and psychopathology composite scores. For each participant, a composite psychopathology score was calculated by multiplying their self-report scores with their CCA-derived loadings, and a composite model parameter score was calculated by multiplying their parameter fits and corresponding loadings. Each dot represents an individual participant. Black line: group-level posterior slope. Gray shading: standard error. (D) Significance testing. p-values for the observed CCA correlation (circle) were computed against empirical null distributions (shaded area and bar showing 95% high-density interval). The top null distribution was obtained by shuffling the model parameter fits (top) relative to the psychopathology scores (top), the middle null distribution was obtained by shuffling only the elasticity prior parameter (γelasticity) while keeping all other parameters and self-report scores matched (middle), and the bottom null distribution was obtained by shuffling only the Sense of Agency score (SOA) while keeping all other parameters and self-report scores matched (bottom).

Revised version: Model parameters fitted to each participant’s task behavior were used as predictors of opt-in (left; multiple logistic regressions) and extra ticket purchases (right; multiple probit regressions) separately for low controllability planets where opting out is optimal (red), high elastic controllability planets where purchasing 3 tickets is optimal (blue), and high inelastic controllability planets where purchasing 1 ticket is optimal (green). In a separate analysis, participants’ composite psychopathology scores (see panel C) were similarly used to predict opt-in (left; logistic regression) and extra ticket purchases (right; probit regression). Bars represent regression coefficients and 95% CIs. (B) Model parameters and psychopathology. Loadings from a Canonical Correlation Analysis (CCA) between model parameters and self-report psychopathology measures. Self-report measures were scored such that higher values reflect higher levels of psychopathology. Loadings are shown for the Initial (dark blue), Replication (light blue), and combined (gray bars) datasets. The magnitude of the loading reflects the degree to which each measure contributes to the correlation between parameters and self-reports. (C) Model parameters and psychopathology composite scores. For each participant, a composite psychopathology score was calculated by multiplying their self-report scores with their CCA-derived loadings, and a composite model parameter score was calculated by multiplying their parameter fits and corresponding loadings. Each dot represents an individual participant. Black line: group-level posterior slope. Gray shading: standard error. (D) Significance testing. p-values for the observed CCA correlation (circle) were computed against empirical null distributions (shaded area and bar showing 95% high-density interval). The top null distribution was obtained by shuffling the model parameter fits (top) relative to the psychopathology scores (top), the middle null distribution was obtained by shuffling only the elasticity prior parameter (𝛾elasticity) while keeping all other parameters and self-report scores matched (middle), and the bottom null distribution was obtained by shuffling only the Sense of Agency score (SOA) while keeping all other parameters and self-report scores matched (bottom).

Earlier version: In contrast, a higher elasticity bias was associated with the purchasing of more extra tickets on planets with inelastic, as compared to elastic, controllability (Initial study: β = .18, 95% CI = [0.03, 0.34], p = .02; Replication study: β = .15, 95%CI = [0.02 0.28], p = .02; linear regression). Thus, a prior assumption that control is likely elastic was primarily reflected in a needless investment of extra resources in environments where control could be obtained more cheaply.

Revised version: In contrast, a higher elasticity bias was associated with the purchasing of more extra tickets on planets with inelastic, as compared to elastic, controllability (Initial study: 𝛽 = .18, 95% CI = [0.03, 0.34], p = .02; Replication study: 𝛽 = .15, 95%CI = [0.02 0.28], p = .02; linear regression). Thus, a prior assumption that control is likely elastic available (operationalized by 𝛾elasticity) was primarily reflected in a needless investment of extra resources in environments where control could be obtained more cheaply.

Other parameters showed associations consistent with their roles in the model. Specifically, the inverse temperature parameter (β) was associated with a more optimal allocation of resources, all three baseline preference parameters (α1, α2, α3) were associated with opting in, and α2 and α3, in particular, were also associated with the purchasing of more extra tickets (Figure 6A).

In sum, the model parameters captured meaningful individual differences in how participants allocated their resources across environments, with the controllability parameter primarily explaining variance in resource allocation in uncontrollable environments, and the elasticity parameter primarily explaining variance in resource allocation in environments where control was inelastic.

Individual elasticity biases are associated with psychopathology

To examine whether the individual biases in controllability and elasticity inference have psychopathological ramifications, we assayed participants on a range of self-report measures of psychopathologies previously linked to a distorted sense of control (see Methods). Examining the direct correlations between model parameters and psychopathology measures (reported in Supplementary Figure 3) does not account for the substantial variance that is typically shared among different forms of psychopathology. For this reason, we instead used a canonical correlation analysis (CCA) to identify particular dimensions within the parameter and psychopathology spaces that most strongly correlate with one another. The CCA loadings indicate the contribution of each parameter of psychopathology score to the overall canonical correlation. We assessed statistical significance of the canonical correlation against a null distribution obtained by permuting the model parameter relative to the psychopathology scores. This showed that the correlation was significant in both the Initial (p=.01) and the Replication (p=.02) studies.

To obtain more robust estimates of the CCA loadings, we next ran a canonical correlation analysis on both datasets combined (Figure 6C, ρ=.33, p=.004). On the side of model parameters, the elasticity prior parameter (𝛾elasticity) received the highest loading. By comparison, the controllability prior parameter received a lower loading with the opposite sign. A further permutation test showed that the contribution of the elasticity prior parameter to the parameter-psychopathology correlation was itself significant (p=.04; Figure 6D).

Loadings on the side of psychopathology were dominated by an impaired sense of agency (SOA; contribution to canonical correlation: p=.03, Figure 6D, bottom plot), along with obsessive compulsive symptoms (OCD), and social anxiety (LSAS) – all symptoms that have been linked to an impaired sense of control25–27. Conversely, depression scores (SDS) received a negative loading, consistent with evidence of a lower willingness to expend one’s resources in depression23,28–29. This psychopathological profile also manifested in raw task measures, since, like the elasticity prior parameter, it too was associated with the purchasing of more tickets on planets where controllability was less elastic (Initial: 𝛽 = .19, 95% CI = [. 06, .32], p=.004; Replication: 𝛽 = .13, 95% CI = [. 01, .25], p = .03; linear regression; Figure 6A). Thus, a distinct psychopathological profile involving a distorted sense of agency was associated with a bias in elasticity inference.

Discussion

Earlier version: Across two pre-registered studies, we demonstrate that humans infer the elasticity of control over their present environment, and this informs their decisions of whether and how to act.

Revised version: Across two pre-registered studies, we demonstrate that humans infer the elasticity of control over their present environment, and this informs their decisions of whether and how to act. Accordingly, we found that a bias to overestimate the elasticity of control leads to a mismanagement of resources, and is associated with elevated psychopathology involving an impaired sense of agency.

Removed in the revised version: Accordingly, we found that a bias to overestimate the elasticity of control leads to a mismanagement of resources, and is associated with elevated psychopathology involving an impaired sense of agency.

Unlike prior work that treated controllability and its elasticity as coupled, our studies dissociated elastic and inelastic controllability, revealing that individuals who typically assume that taking action is worthwhile (overall controllability) do not necessarily assume that investing more resources in taking action is worthwhile (the elasticity of control). In fact, biases regarding controllability and its elasticity were oppositely related to psychopathology, revealing a psychopathological profile that combines underestimation of control with overestimation of the degree to which investing more resources is necessary to obtain control.

The observed elasticity biases across individuals are particularly striking given the simplicity of the experimental task in comparison to real-life scenarios. Unlike the task, where successful boarding directly determined a trip’s outcome, in many real situations many actions may contribute to a single outcome. In this regard, the present work introduces a new kind of credit-assignment problem30, which involves identifying not only which actions influenced the outcome but also which of these actions was made more effective by investing more resources in its execution. Moreover, whereas in the task elasticity increased linearly with the number of tickets, resource investment in real life often exhibits non-linear effects. Specifically, further increases in resource investment often result in diminishing returns, and beyond a certain point, even worse outcomes (e.g. hiring too many employees, making the work less efficient). Additionally, real life typically doesn’t offer the streamlined recurrence of homogenized experiences that makes learning easier in experimental tasks, nor are people systematically instructed and trained about elastic and inelastic control in each environment. These complexities introduce substantial additional uncertainty into inferences of elasticity in naturalistic settings, thus allowing more room for prior biases to exert their influences. The elasticity biases observed in the present studies are therefore likely to be amplified in real-life behavior. Future research should examine how these complexities affect judgments about the elasticity of control to better understand how people allocate resources in real-life.

Earlier version: Our finding that elasticity biases are associated with agency-related psychopathologies raises interesting questions for future research on the causal nature of this association. One possibility is that biased elasticity inferences distort one’s sense of agency by leading to maladaptive resource allocation that results in consistently suboptimal outcomes. Another possibility is that people who already have a reduced sense of agency, attempt to compensate for this feeling by investing more resources despite minimal benefit, thus demonstrating an elasticity bias.

Revised version: Our finding that elasticity biases are associated with agency-related psychopathologies raises interesting questions for future research on the causal nature of this association. One possibility is that biased elasticity inferences distort one’s sense of agency by leading to maladaptive resource allocation that results in consistently suboptimal outcomes. Another possibility is that people who already have a reduced sense of agency, attempt to compensate for this feeling by investing more resources despite minimal benefit, thus demonstrating an elasticity bias. Importantly, these possibilities are not mutually exclusive. That is, elasticity biases and a distorted sense of agency could reinforce each other in a cyclical manner. Longitudinal studies tracking or manipulating elasticity beliefs while monitoring changes in real-life resource investment and subjective sense of agency, could resolve how elasticity biases and psychopathology influence each other.

Removed in the revised version: Importantly, these possibilities are not mutually exclusive. That is, elasticity biases and a distorted sense of agency could reinforce each other in a cyclical manner. Longitudinal studies tracking or manipulating elasticity beliefs while monitoring changes in real-life resource investment and subjective sense of agency, could resolve how elasticity biases and psychopathology influence each other.

Removed in the revised version: Whereas our pre-registered CCA effectively identified associations between task parameters and a psychopathological profile, this analysis method does not directly reveal relationships between individual variables. Auxiliary analyses confirmed significant contributions of both elasticity bias and sense of agency to the observed canonical correlation, but the contribution of other measures remains to be determined by future work. Such work could employ other established measures of agency, including both behavioral indices and subjective self-reports, to better understand how these constructs relate across different contexts and populations.

Earlier version: Our operationalization of controllability aligns with previous work that defined controllability as the degree to which actions influence the probability of obtaining a reward5–6,6. Other notable works defined controllability as the degree to which actions predict subsequent states, independent of reward1–3 Here we adhered to the former definition due to our present focus on resource elasticity, since individuals are unlikely to invest resources to gain control if control does not yield more reward. That said, we followed the latter work in ensuring that controllability is not confounded with predictability2, in the sense that in planets with high controllability, outcomes were not necessarily more predictable than in planets with low controllability (See Methods). This enabled us to specifically assess individual inferences concerning controllability and its resource elasticity.

Revised version: Whereas our pre-registered CCA effectively identified associations between task parameters and a psychopathological profile, this analysis method does not directly reveal relationships between individual variables. Auxiliary analyses confirmed significant contributions of both elasticity bias and sense of agency to the observed canonical correlation, but the contribution of other measures remains to be determined by future work. Such work could employ other established measures of agency, including both behavioral indices and subjective self-reports, to better understand how these constructs relate across different contexts and populations. Our operationalization of controllability aligns with previous work that defined controllability as the degree to which actions influence the probability of obtaining a reward5–6,6. Other notable works defined controllability as the degree to which actions predict subsequent states, independent of reward1–3 Here we adhered to the former definition due to our present focus on resource elasticity, since individuals are unlikely to invest resources to gain control if control does not yield more reward. That said, we followed the latter work in ensuring that controllability is not confounded with predictability2, in the sense that in planets with high controllability, outcomes were not necessarily more predictable than in planets with low controllability (See Methods). This enabled us to specifically assess individual inferences concerning controllability and its resource elasticity.

Future research could explore alternative models for implementing elasticity inference that extend beyond our current paradigm. First, further investigation is warranted concerning how uncertainty about controllability and its elasticity interact. In the present study, we minimized individual differences in the estimation of maximum available control by providing participants with three free tickets at their initial visits to each planet. We made this design choice to isolate differences in the estimation of elasticity, as opposed to maximum controllability. To study how these two types of estimations interact, future work could benefit from modifying this aspect of our experimental design.

Earlier version: Second, future models could enable generalization to levels of resource investment not previously experienced. For example, controllability and its elasticity could be jointly estimated via function approximation that considers control as a function of invested resources. Although our implementation of this model did not fit participants’ choices well (see Methods), other modeling assumptions or experimental designs may offer a better test of this idea.

Revised version: Second, future models could enable generalization to levels of resource investment not previously experienced. For example, controllability and its elasticity could be jointly estimated via function approximation that considers control as a function of invested resources. Although our implementation of this model did not fit participants’ choices well (see Methods), other modeling assumptions drawn from human functional learning31 or experimental designs with continuous action spaces may offer a better test of this idea.

Another interesting possibility is that individual elasticity biases vary across different resource types (e.g., money, time, effort). For instance, a given individual may assume that controllability tends to be highly elastic to money but inelastic to effort. Although the task incorporated multiple resource types (money, time, and attentional effort), the results may differ depending on the type of resources on which the participant focuses. Future studies could explore this possibility by developing tasks that separately manipulate elasticity with respect to different resource types. This would clarify whether elasticity biases are domain-specific or domain-general, and thus elucidate their impact on everyday decision-making.

Methods

Pre-registration

Both the Initial and Replication studies were pre-registered prior to data collection. The pre-registrations included hypotheses, sample sizes, exclusion criteria, task design, analysis plans (regressions, CCA), and computational models. Detailed pre-registration documents are available at https://aspredicted.org/Q7M_BHR and https://aspredicted.org/CHW_12H.

Participants

Earlier version: We used the Prolific online platform (prolific.com) to recruit participants (Initial study N = 264, Replication study N = 250). Participants were selected based on their demonstrated high engagement levels in prior studies, fluency in English, and absence of psychiatric diagnoses. Due to cultural differences in the perception of controllability31 participants were recruited from English speaking western countries. The average age of participants across both datasets was 33 years old (SE =.6), and 38% were female.

Revised version: We used the Prolific online platform (prolific.com) to recruit participants (Initial study N = 264, Replication study N = 250). Participants were selected based on their demonstrated high engagement levels in prior studies, fluency in English, and absence of psychiatric diagnoses. Due to cultural differences in the perception of controllability32 participants were recruited from English speaking western countries. The average age of participants across both datasets was 33 years old (SE =.6), and 38% were female.

The sample size was determined using a power analysis as described below. Exclusion criteria were applied to ensure data quality and participant engagement. Specifically, participants were excluded if they demonstrated inattentiveness (defined as more than 8 mistakes on comprehension quizzes during the instruction phase), displayed less than 90% accuracy in selecting the correct vehicle corresponding to its pre-learned destination (mean error rate = 3%, range [0-10%]; indicating a lack of task comprehension), or chose not to purchase any tickets in 90% or more of trials across all experimental conditions (suggesting a lack of engagement with the task). These criteria led to the exclusion of 22 participants (8.3% exclusion rate) from the initial study and 23 participants (8.9% exclusion rate) from the replication study. Participants gave written informed consent before taking part of the study, which was approved by the university’s ethics review board (approval number: 2022-06213). Participants were paid at a rate of 9 euros per hour.

Experimental Task

Participants were informed about the location of a treasure (150 coins, worth 15¢ of bonus monetary compensation) and could exercise control by boarding the train to reach the house, or the plane to reach the mountain. (see main text and Figure 2). Forgoing or failing to board resulted in walking to the nearest location which happened to be where the treasure was located 20% of the time.

Removed in the revised version: The first ticket (cost = 40 coins) allowed participants to select their desired vehicle, but it only stopped for them at the platform some percentage of the time, corresponding to the inelastic controllability of the planet. Participants could purchase up to 2 additional tickets (cost = 20 coins each) to attempt boarding the moving vehicle after it had left the platform. This improved the probability of boarding in correspondence with the elastic controllability of the planet.

Earlier version: When participants purchased a second or third ticket, the chosen vehicle appeared moving from left to right across the screen, and participants attempted to board it by pressing the spacebar when it reached the center of the screen. The precision of each boarding attempt was quantified as the absolute distance of the vehicle from the screen’s center at the time of the press. Unless elastic control was below 15%, boarding attempts with precision in the top 15 percentile relative to the participant’s prior performance automatically succeeded, and unless elastic control was above 85%, those in the bottom 15 percentile automatically failed. For the remaining attempts, the outcome was probabilistically determined such that the overall probability of successfully boarding matched the elasticity of control on that planet. To verify that participants’ pressing ability did not bias their preference for purchasing extra tickets, we ran a mixed probit regression of the number of extra tickets purchased on participants’ average press precision. We found no significant effect (p = .45). When participants purchased multiple tickets, they made all boarding attempts in sequence without intermediate feedback, only learning whether they successfully boarded upon reaching their final destination. This served two purposes. First, to ensure that participants would consistently need to invest not only more money but also more effort and time in planets with high elastic controllability. Second, to ensure that results could potentially generalize to the many real-world situations where the amount of invested effort has to be determined prior to seeing any outcome (e.g., preparing for an exam or a job interview).

Revised version: The first ticket (cost = 40 coins) allowed participants to select their desired vehicle, but it only stopped for them at the platform some percentage of the time, corresponding to the inelastic controllability of the planet. Participants could purchase up to 2 additional tickets (cost = 20 coins each) to attempt boarding the moving vehicle after it had left the platform. This improved the probability of boarding in correspondence with the elastic controllability of the planet. When participants purchased a second or third ticket, the chosen vehicle appeared moving from left to right across the screen, and participants attempted to board it by pressing the spacebar when it reached the center of the screen. The precision of each boarding attempt was quantified as the absolute distance of the vehicle from the screen’s center at the time of the press. Unless elastic control was below 15%, boarding attempts with precision in the top 15 percentile relative to the participant’s prior performance automatically succeeded, and unless elastic control was above 85%, those in the bottom 15 percentile automatically failed. For the remaining attempts, the outcome was probabilistically determined such that the overall probability of successfully boarding matched the elasticity of control on that planet. To verify that participants’ pressing ability did not bias their preference for purchasing extra tickets, we ran a mixed probit regression of the number of extra tickets purchased on participants’ average press precision. We found no significant effect (p = .45). When participants purchased multiple tickets, they made all boarding attempts in sequence without intermediate feedback, only learning whether they successfully boarded upon reaching their final destination. This served two purposes. First, to ensure that participants would consistently need to invest not only more money but also more effort and time in planets with high elastic controllability. Second, to ensure that results could potentially generalize to the many real-world situations where the amount of invested effort has to be determined prior to seeing any outcome (e.g., preparing for an exam or a job interview).

At the end of each trip, participants were shown where they arrived, allowing them to infer whether they successfully boarded their vehicle. Since this feedback is only interpretable if participants knew each vehicle’s destination, we only included participants in subsequent analyses if they chose the vehicle corresponding to the destination of the treasure with at least 90% accuracy (Mean 97% ±2%, see first exclusion criteria). The average trial duration was approximately 11 seconds (SE = 0.05), with each additional ticket extending this duration by approximately 0.5 seconds (SE = 0.04), increasing the time investment required from participants.

To allow sufficient learning of each planet’s elasticity and controllability, participants took 30 consecutive trips on each planet. This trip count was determined based on pilot data (N=19) by analyzing when participants’ strategies stabilized. Specifically, we examined the standard deviation in ticket choices over the final 5 trips and found stabilization after around 15 trips (STD = 0.33 compared to .63 in the first 15 trips). Based on this, we provided 15 initial trips per planet for learning, and analyzed ticket choices starting from trip 16 onward. This approach followed our pre-registered plan.

To homogenize learning experience across participants, we gave participants three free tickets for the first five visits to every planet. However, only four of those trips were informative, since on one of every five trips (randomly placed), the treasure was located in the nearby location, making boarding inconsequential for reaching it.

Earlier version: Points accumulated across all planets throughout the session, with participants explicitly motivated to maximize their total points as this directly determined their monetary bonus payment. To ensure that accumulated gains did not lead participants to adopt a simple heuristic strategy of always purchasing 3 tickets, we conducted a mixed probit regression examining whether the number of accumulated coins influenced participants’ decisions to purchase extra tickets. We did not find such an effect (βcoin accumulated = .01 p = .87), ruling out the potential strategy shift.

Revised version: Points accumulated across all planets throughout the session, with participants explicitly motivated to maximize their total points as this directly determined their monetary bonus payment. To ensure that accumulated gains did not lead participants to adopt a simple heuristic strategy of always purchasing 3 tickets, we conducted a mixed probit regression examining whether the number of accumulated coins influenced participants’ decisions to purchase extra tickets. We did not find such an effect (βcoins accumulated = .01 p = .87), ruling out the potential strategy shift.

Planet types

To study elasticity learning across all possible degrees of controllability and elasticity, 66 unique planet conditions were generated by systematically varying inelastic controllability (the probability of the chosen vehicle stopping at the platform if purchasing a single ticket) and elasticity (the added probability of boarding from purchasing 2 extra tickets) from 0 to 100% in 10% increments. The treasure value (150 coins) and ticket costs (40 coins for the first ticket, 20 coins for extra tickets) were chosen so that the optimal strategy, as determined by expected value calculations (equation 11), was to purchase 0 tickets on 22 planets (low controllability condition), 1 ticket on a different set of 22 planets (high inelastic controllability condition), and 3 tickets on the remaining 22 planets (high elastic controllability condition). Participants were randomly assigned one planet from each condition, while ensuring they could earn roughly the same total number of coins if they made optimal choices (Initial study= 45±.3, Replication study 46 ± .3). To address potential confounds between controllability and predictability2, we designed the experiment such that high-controllability planets (where purchasing 1 or 3 tickets was optimal) were not inherently more predictable than low-controllability planets. We quantified predictability as the entropy of the probability distribution determining whether the participant will or will not arrive at the rewarded location. For 60% of participants, purchasing tickets on high-controllability planets decreased entropy (i.e., increased predictability) compared to low-controllability planets, whereas for the other 40% it increased entropy (i.e., decreased predictability). To increase sensitivity to individual differences, all participants first encountered a planet where purchasing any number of tickets (0, 1, 2, or 3) yielded the same expected value (EV = 30 coins).

Initial training

Earlier version: Prior to engaging in the task, participants completed 90 practice trials. These practice trials were divided equally across planets with high inelastic controllability, high elastic controllability, and low controllability. To reinforce participants’ understanding of how elasticity and controllability were manifested in each planet, they were informed of the planet type they had visited after every 15 trips. Specifically, for high inelastic controllability planets, the message read: “ Your ride actually would stop for you with 1 ticketn So purchasing extra tickets, since they do cost money, is a WASTE.” For low controllability planets: “Doesn’t seem like the vehicle stops for you nor does purchasing extra tickets help.” Lastly, for high elastic controllability planets: “Hopefully by now it’s clear that only by purchasing 3 tickets (LOADING AREA) are you consistently successful in catching your ride.”.

Revised version: Prior to engaging in the task, participants completed 90 practice trials. These practice trials were divided equally across planets with high inelastic controllability, high elastic controllability, and low controllability. To reinforce participants’ understanding of how elasticity and controllability were manifested in each planet, they were informed of the planet type they had visited after every 15 trips. Specifically, for high inelastic controllability planets, the message read: “Your ride actually would stop for you with 1 ticket! So purchasing extra tickets, since they do cost money, is a WASTE.” For low controllability planets: “Doesn’t seem like the vehicle stops for you nor does purchasing extra tickets help.” Lastly, for high elastic controllability planets: “Hopefully by now it’s clear that only by purchasing 3 tickets (LOADING AREA) are you consistently successful in catching your ride.”.

Power analysis

Earlier version: The sample sizes for both the initial and the replication studies were pre-registered before data collection commenced. For the initial study, we derived the sample size of 264 participants using a power analysis aiming to estimate with high accuracy (credible interval width < 0.1) the likelihood that participants will opt-in for any combination of elastic and inelastic control (66 total combinations). For this purpose, we used Kruschke’s (2014) procedure for Bayesian Power Analysis32. Based on pilot data (N=19), we fitted a Bayesian Beta Binomial model to the proportion of opt-ins in the planet with no optimal strategy (and therefore the highest variability in behavior would be expected), and found that 12 participants were required for the credible interval of the binomial proportion parameter to be lower than 0.1. Since each participant completes 3 combinations of elastic and inelastic control, we multiplied the required sample size by the total number of planets in each condition (22) which results in an estimate of n=264.

Revised version: The sample sizes for both the initial and the replication studies were pre-registered before data collection commenced. For the initial study, we derived the sample size of 264 participants using a power analysis aiming to estimate with high accuracy (credible interval width < 0.1) the likelihood that participants will opt-in for any combination of elastic and inelastic control (66 total combinations). For this purpose, we used Kruschke’s (2014) procedure for Bayesian Power Analysis33. Based on pilot data (N=19), we fitted a Bayesian Beta Binomial model to the proportion of opt-ins in the planet with no optimal strategy (and therefore the highest variability in behavior would be expected), and found that 12 participants were required for the credible interval of the binomial proportion parameter to be lower than 0.1. Since each participant completes 3 combinations of elastic and inelastic control, we multiplied the required sample size by the total number of planets in each condition (22) which results in an estimate of n=264.

Earlier version: To examine whether this sample size is sufficient to detect significant effects of elastic and inelastic controllability on opting in (Figure 3A,C), and the purchasing of extra tickets (Figure 3B,C), we used a bootstrap procedure whereby we sampled 264 participants from the pilot data and ran the regression analyses on them. We repeated this procedure for 50 iterations, each time recording βelastic, β*nelastic, and their p-values. For analysis 1, we found that both βelastic and βinelastic were significantly above 0 in all iterations, and for analysis 2, βelastic was significantly above 0 in all iterations, thereby giving us >98% power to detect all of our effects of interest. Confidence intervals were derived from a binomial test in which a positive significant coefficient counted as a success.

Revised version: To examine whether this sample size is sufficient to detect significant effects of elastic and inelastic controllability on opting in (Figure 3A,C), and the purchasing of extra tickets (Figure 3B,C), we used a bootstrap procedure whereby we sampled 264 participants from the pilot data and ran the regression analyses on them. We repeated this procedure for 50 iterations, each time recording 𝛽elastic, 𝛽inelastic, and their p-values. For analysis 1, we found that both 𝛽elastic 𝑎𝑛𝑑 inelastic were significantly above 0 in all iterations, and for analysis 2, 𝛽elastic was significantly above 0 in all iterations, thereby giving us >98% power to detect all of our effects of interest. Confidence intervals were derived from a binomial test in which a positive significant coefficient counted as a success.

For the replication study, we derived the sample size of 250 participants using a power analysis aiming for 95% power to detect a significant (p <.05) canonical correlation between the model-derived parameters and questionnaire scores (Figure 6B). For this purpose, we sampled with replacement from the initial dataset (n=264) which had showed a significant canonical correlation (p=.01; Figure 6B), and calculated the p-value of a CCA in the subsample. We repeated this procedure for 100 iterations, each time recording the calculated p-value. To ensure our sample size was sufficient to detect significant effects in the regression models predicting opt-in and extra ticket purchases from the composite psychopathology score with over 95% power (Figure 6C) we conducted another bootstrap analysis. We sampled 250 participants with replacement each time calculating the composite score and ran the regression analyses on that score, repeating this procedure 100 times. We found that 𝛽composite was significant (p<.05) in 95% of iterations (95% power).

Regression analyses

Earlier version: Mixed model regression analyses were performed using the “ordinal” package which fits cumulative mixed models via Laplace approximations33. The model predicts the likelihood of purchasing tickets (logistic model) and how many extra tickets are purchased (probit model) on the last 15 trips on each planet. Predictors included fixed and random effects for elastic and inelastic controllability and fixed and random intercepts per participant and planet number, with the latter included to account for potential learning or fatigue effects. Statistical significance was assessed using the Wald Z Test34 with normality assessed using the Shapiro-Wilk Test35.

Revised version: Mixed model regression analyses were performed using the “ordinal” package which fits cumulative mixed models via Laplace approximations34. The model predicts the likelihood of purchasing tickets (logistic model) and how many extra tickets are purchased (probit model) on the last 15 trips on each planet. Predictors included fixed and random effects for elastic and inelastic controllability and fixed and random intercepts per participant and planet number, with the latter included to account for potential learning or fatigue effects. Statistical significance was assessed using the Wald Z Test35 with normality assessed using the Shapiro-Wilk Test36.

Computational Modeling

To assess whether participants inferred elasticity or merely learned the control afforded by each course of action separately, we fit two competing reinforcement-learning models (‘controllability’ and ‘elastic controllability’). Both models use beta distributions defined by parameters (a, b), to form controllability beliefs for each ticket quantity n (n ∈ {1, 2, 3}), but only the elastic controllability model leverages the dependencies that exist between different levels of resource investment. Below, we expand on the computations underlying both models.

Controllability model

The controllability model accumulates the number of times that each ticket quantity n led to successful or unsuccessful boarding: Where O=1 if boarding was successful and O=0 if not. The probability of boarding with 1,2 or 3 tickets on planet s is thus simply the expected value of that particular beta distribution: Here and in all subsequent models, learning only takes place when the treasure is not located on the nearby planet, which makes it possible to identify boarding success.

Elastic controllability model

Earlier version: The elastic controllability model explains ticket choices based on latent beliefs about the controllability and elasticity of each planet (see Figure 4). These beliefs are represented by three beta distributions, each defined by two parameters (a, b). The expected value of one beta distribution (defined by acontrol, bcontrol) represents the belief that boarding is possible (controllability) with any number of tickets whereas the expected value of the other two beta distributions represent the belief that successful boarding is added by purchasing at least one (aelastic≥1,, belastic2) or specifically two (aelastic2, belastic2), extra tickets (elasticity).

Revised version: The elastic controllability model explains ticket choices based on latent beliefs about the controllability and elasticity of each planet (see Figure 4). These beliefs are represented by three beta distributions, each defined by two parameters (𝑎, 𝑏). The expected value of one beta distribution (defined by 𝑎control, 𝑏control) represents the belief that boarding is possible (controllability) with any number of tickets whereas the expected value of the other two beta distributions represent the belief that successful boarding is added by purchasing at least one (𝑎elastic≥1, 𝑏elastic2) or specifically two (𝑎elastic2, 𝑏elastic2), extra tickets (elasticity).

Earlier version: Thus, the parameter acontrol accumulates the number of successful boarding attempts with any number of tickets, whereas bcontrol accumulates the number of failed boarding attempts despite purchasing the maximum number of tickets (3 tickets): A failure to board with fewer tickets doesn’t reduce the maximum controllability estimate but instead suggests that more tickets or the maximum number of tickets may be needed to obtain control: Conversely, successfully boarding with fewer tickets is evidence that obtaining controllability does not require maximal or any extra tickets: To account for our finding that failure with 3 tickets made participants favor 3, over 1 and 2, tickets, we introduced a modified elastic controllability* model, wherein purchasing extra tickets is also favored upon receiving evidence of low controllability (loss with 3 tickets). This effect was modulated by a free parameter κ which reflects a tendency to prioritize resolving uncertainty about whether control is at all possible by investing maximum resources: The estimated controllability and elasticity, as characterized by the expected value of the three beta distributions (Ctotal control, Cmore, Cfull), is used to derive the probability of boarding the selected vehicle with 1, 2, or 3 tickets. Thus, the estimated probability of boarding with 1 ticket is an estimate of inelastic control, that is, the percentage of the total control Ctotal control that does not require investing more resources: For 2 tickets, the added probability of boarding is an estimate of what percentage of total control is elastic (Cmore) but does not require investing maximal resources (1 − Cfull): Finally, since 3 tickets is the maximum resource investment, the estimated probability of boarding with 3 tickets is equivalent to an estimate of total attainable control.

Revised version: Thus, the parameter 𝑎control accumulates the number of successful boarding attempts with any number of tickets, whereas 𝑏control accumulates the number of failed boarding attempts despite purchasing the maximum number of tickets (3 tickets): A failure to board with fewer tickets doesn’t reduce the maximum controllability estimate but instead suggests that more tickets or the maximum number of tickets may be needed to obtain control: Conversely, successfully boarding with fewer tickets is evidence that obtaining controllability does not require maximal or any extra tickets: To account for our finding that failure with 3 tickets made participants favor 3, over 1 and 2, tickets, we introduced a modified elastic controllability* model, wherein purchasing extra tickets is also favored upon receiving evidence of low controllability (loss with 3 tickets). This effect was modulated by a free parameter 𝜅 which reflects a tendency to prioritize resolving uncertainty about whether control is at all possible by investing maximum resources: The estimated controllability and elasticity, as characterized by the expected value of the three beta distributions (𝐶total control, 𝐶more, 𝐶full), is used to derive the probability of boarding the selected vehicle with 1, 2, or 3 tickets. Thus, the estimated probability of boarding with 1 ticket is an estimate of inelastic control, that is, the percentage of the total control 𝐶total control that does not require investing more resources: For 2 tickets, the added probability of boarding is an estimate of what percentage of total control is elastic (Cmore) but does not require investing maximal resources (1 − Cfull): Finally, since 3 tickets is the maximum resource investment, the estimated probability of boarding with 3 tickets is equivalent to an estimate of total attainable control.

Models’ resource investment

Because participants win 20% of the time even when they fail to board their selected ride, we calculate the probability of arriving at the treasure with n tickets as: Then, taking the possible reward (R = 150) and ticket costs (C = {40,20,20}) into account, the expected value for purchasing n tickets (n = {1,2,3}) tickets is: One exception to this rule is the expected value of opting out, V(0), which is similar across all planets, and equals the potential reward (150) times the probability that walking will lead to the reward (0.2).

Earlier version: The probability that a participant will purchase 0, 1, 2, or 3 tickets is calculated using the SoftMax function such that: where β is an inverse temperature parameter, α1−3 are parameters that represent participants’ fixed tendencies to purchase 1, 2, or 3 tickets respectively, and ρ is a perseveration parameter that represents participant tendency to purchase the same number of tickets as in the last trial (nt−1). A0 is set to 0 to maintain parameter identifiability.

Revised version: The probability that a participant will purchase 0, 1, 2, or 3 tickets is calculated using the SoftMax function such that: where β is an inverse temperature parameter, α1–3 are parameters that represent participants’ fixed tendencies to purchase 1, 2, or 3 tickets respectively, and ρ is a perseveration parameter that represents participant tendency to purchase the same number of tickets as in the last trial (nt–1). ΑO is set to 0 to maintain parameter identifiability.

We compared the aforementioned models to a null model wherein participants have distinct preferences among different number of tickets (α1–3) and perseveration (ρ) but no latent controllability estimates (𝛽 = 0 in equation 12). We also evaluated variants of both models that replaced the SoftMax-based exploration (governed by inverse temperature β) with a decaying epsilon-greedy approach. This epsilon-greedy variant, which begins with a fixed exploration rate that decreases with the number of trials59, performed significantly worse at explaining participants’ choices compared to the SoftMax-based models.

Prior controllability and elasticity beliefs

To capture prior biases that planets are controllable and elastic, we introduced parameters 𝛾controllability and 𝛾elasticity, each computed by multiplying the direction (λ – 0.5) and strength (ɛ) of individuals’ prior belief. λcontrollability and λelasticity range between 0 and 1, with values above 0.5 indicating a bias towards high controllability or elasticity, and values below 0.5 indicating a bias towards low controllability or elasticity. ɛcontrollability and ɛelasticity are positively valued parameters capturing confidence in the bias. Parameter recovery analyses confirmed both good recoverability (see S2 Table) and low confusion between bias direction and strength (ɛcontrollability→λcontrollability = −.07, λcontrollability→ɛcontrollability = −.16, ɛelasticity→λelasticity = .15, λelasticity→ɛelasticity = .04), ensuring that observed biases and their relation to psychopathology do not merely reflect slower learning (Supplementary Figure 4), which can result from changes in bias strength but not direction.

Earlier version: For each new planet participants encountered, these parameters were used to initialize the beta distributions representing participants’ beliefs: a controllability is multiplied by 2 so that in the absence of a bias, it would have equal preference for purchasing 0 to 3 tickets (this requires that ).

Revised version: For each new planet participants encountered, these parameters were used to initialize the beta distributions representing participants’ beliefs: 𝑎 controllability is multiplied by 2 so that in the absence of a bias, it would have equal preference for purchasing 0 to 3 tickets (this requires that C controllability = , C more,full = ).

For consistency, in the ‘controllability model’, the controllability estimates for each ticket quantity are initialized for each ticket quantity n, where n={1,2,3}, as: Here, to ensure no preference between different ticket quantities in the absence of a bias, 𝑎1 is multiplied by 2 and 𝑎2 is multiplied by .

Testing alternative models

To further validate our elastic controllability model, we tested it against several theoretically motivated alternatives that could potentially explain participants’ behavior.

To ascertain that participants were truly learning latent estimates of controllability rather than simpler associations, we conducted two complementary analyses.

First, we implemented a simple Q-learning model that directly maps ticket quantities to expected values based on reward prediction errors, without representing latent controllability. This associative model performed substantially worse than even our simple controllability model (log Bayes Factor ≥ 1854 on the combined datasets). Second, we fitted a variant of the elastic controllability model that compared learning from control-related versus chance outcomes via separate parameters (instead of assuming no learning from chance outcomes). Chance outcomes were observed by participants in the 20% of trials where reward and control were decoupled, in the sense that participants reached the treasure regardless of whether they boarded their vehicle of choice. Results showed that participants learned considerably more from control-related, as compared to chance, outcomes (mean learning ratio=1.90, CI= [1.83, 1.97]). Together, these analyses show that participants were forming latent controllability estimates rather than direct action-outcome associations.

Having established that participants indeed learn latent controllability estimates, we next examined our model’s interpretation of the κ parameter. Our model accounts for the observation that failures with 3 tickets made participants more likely to purchase 3 tickets when they opted in. To test whether this behavior resulted from participants not accepting that the planet is uncontrollable, we implemented a variant of the model where κ not only increases the elasticity estimate but also reduces the controllability update (using 𝛽control+(1-𝜅) instead of 𝛽control+1) after failures with 3 tickets. However, implementing this coupling diminished the model’s fit to the data (log Bayes Factor ≥ 48), as compared to allowing both effects to occur independently, indicating that the increase in 3 ticket purchases upon failing with 3 tickets did not result from participants not accepting that controllability is in fact low. Thus, we maintained our original interpretation that failure with 3 tickets increases uncertainty about whether control is possible at all, leading participants who continue to opt in to invest maximum resources to resolve this uncertainty.

Having validated our interpretation of κ, we next examined whether participant behavior would be better characterized as a continuous function approximation rather than the discrete inferences in our model. To test this, we implemented a Bayesian model where participants continuously estimate both controllability and its elasticity as a mixture of three archetypal functions mapping ticket quantities to success probabilities. The flat function provides no control regardless of how many tickets are purchased (corresponding to low controllability). The step function provides full control as long as at least one ticket is purchased (inelastic controllability). The linear function linearly increases control with the number of extra tickets (i.e., 0%, 50%, and 100% control for 1, 2, and 3 tickets, respectively; elastic controllability). The model computes the likelihood that each of the functions produced each new observation, and accordingly updates its beliefs. Using these beliefs, the model estimates the probability of success for purchasing each number of tickets, allowing participants to weigh expected control against increasing ticket costs. Despite its theoretical advantages for generalization to different ticket quantities, this continuous function approximation model performed significantly worse than the elastic controllability model (log Bayes Factor > 4100 on combined datasets), suggesting that participants did not assume that control increases linearly with resource investment.

Model fitting

Earlier version: Reinforcement-learning models were fitted to all choices made by participants via an expectation maximization approach used in previous work36–38. We generated random parameter settings drawn from predefined group-level distributions and calculated the likelihood of observing the participants’ choices given each parameter setting. We then approximated the posterior estimates of the group-level prior distributions for each parameter by resampling the parameter values, weighted by their respective likelihoods, and refitted the data based on the updated priors. This iterative process of resampling and refitting continued until no further improvement to model fit. To determine the best-fitting parameters for each individual participant, we computed a weighted average of the final batch of parameter settings, where each setting was weighted by the likelihood it assigned to the individual participant’s choices.

Revised version: Reinforcement-learning models were fitted to all choices made by participants via an expectation maximization approach used in previous work37–39. We generated random parameter settings drawn from predefined group-level distributions and calculated the likelihood of observing the participants’ choices given each parameter setting. We then approximated the posterior estimates of the group-level prior distributions for each parameter by resampling the parameter values, weighted by their respective likelihoods, and refitted the data based on the updated priors. This iterative process of resampling and refitting continued until no further improvement to model fit. To determine the best-fitting parameters for each individual participant, we computed a weighted average of the final batch of parameter settings, where each setting was weighted by the likelihood it assigned to the individual participant’s choices.

Earlier version: Model parameter β was initialized by sampling from a Gamma distribution (k = 1, θ = 1), λ was initialized by sampling from a Beta distribution (α = 1 and β = 1), ϵ was initialized by sampling from a Lognormal distribution (μ = 0, σ = 1)., and all other parameters (α1−2, ρ, κ) were initialized by sampling from a Normal distribution (μ = 0, σ = 1).

Revised version: Model parameter 𝛽 was initialized by sampling from a Gamma distribution (k = 1, θ = 1), 𝜆 was initialized by sampling from a Beta distribution (α = 1 and β = 1), 𝜖 was initialized by sampling from a Lognormal distribution (μ = 0, σ = 1)., and all other parameters (𝛼1–3, 𝜌, 𝜅) were initialized by sampling from a Normal distribution (μ = 0, σ = 1).

Validation of model fitting procedure

To validate the model fitting procedure, we simulated data using each model and using the model comparison procedure to recover the correct model (Supplementary Table 1). Furthermore, we validated our parameter fits by simulating data using the best fitting parameters for each parameter and then recovering those parameters. The correlations for the parameters of interest were all equal or greater than r = .74 and at least .63 for all other parameters (see Supplementary Table 2).

Model comparison

Earlier version: To compare between models, in terms of how well each model accounted for participants’ choices, we employed the integrated Bayesian Information Criterion (iBIC)39–40. First, we estimated the evidence for each model (L) by calculating the mean likelihood of the model given 200,000 random parameter settings sampled from the fitted group-level prior distributions.

Revised version: To compare between models, in terms of how well each model accounted for participants’ choices, we employed the integrated Bayesian Information Criterion (iBIC)40–41. First, we estimated the evidence for each model (L) by calculating the mean likelihood of the model given 200,000 random parameter settings sampled from the fitted group-level prior distributions. Subsequently, we computed the iBIC by penalizing the model evidence to account for model complexity, following this formula: iBIC = 2 ln L + k ln n, where k represents the number of fitted parameters, and n is the number of subject choices used to compute the likelihood. Lower iBIC values indicate a more parsimonious and better-fitting model.

Removed in the revised version: Subsequently, we computed the iBIC by penalizing the model evidence to account for model complexity, following this formula: iBIC = 2 ln L + k ln n, where k represents the number of fitted parameters, and n is the number of subject choices used to compute the likelihood. Lower iBIC values indicate a more parsimonious and better-fitting model.

Model simulations

To examine the predictions of the ‘controllability’ and ‘elastic controllability’ models, we simulated choices using the best-fitting parameter settings obtained from participants’ actual choices. For generality, planets were randomized in the simulation as they were in the actual experiment, such that a simulated participant did not visit precisely the same planets as the actual participant. Furthermore, to avoid discrepancies between models, each planet was repeated the same number of times for both models. 2000 simulations were performed with each model.

Individual differences

We examined individual biases in elasticity inference using participants’ best-fitting parameters for beliefs about controllability (𝛾controllability) and elasticity (𝛾elasticity). In addition to these focal parameters, we also examined four other parameters that directly influence opt-in decisions and extra ticket purchases: an inverse temperature/noise parameter (β) and three baseline preference parameters (α1, α2, α3).

Earlier version: We then regressed participants’ actual opt-in (logistic regression) and extra ticket purchases (probit regression) in each of the 3 controllability conditions on participants’ parameters.

Revised version: We then regressed participants’ actual opt-in (logistic regression) and extra ticket purchases (probit regression) in each of the 3 controllability conditions on participants’ parameters. Regressions were performed in the lme4 package in R using Rstudio42a,42b. Statistical significance was assessed using the Wald Z Test35 with normality assessed using the Shapiro-Wilk Test36.

Removed in the revised version: Regressions were performed in the lme4 package in R using Rstudio41. Statistical significance was assessed using the Wald Z Test34 with normality assessed using the Shapiro-Wilk Test35.

Self-report measures

Earlier version: To assess trait-level agency beliefs and related psychopathologies, we administered the Sense of Agency Scale (SOA)42. Additionally, we measured dysfunctional attitudes toward problem- solving via the Negative Problem Orientation Questionnaire (NPOQ) 43, and temperamental differences (five domains: hyperthymic, dysthymic, cyclothymic, irritable, and anxious) via the Temperament Evaluation Questionnaire (TEMPS-A)44–45.

Revised version: To assess trait-level agency beliefs and related psychopathologies, we administered the Sense of Agency Scale (SOA)43. Additionally, we measured dysfunctional attitudes toward problem- solving via the Negative Problem Orientation Questionnaire (NPOQ) 44a,44b, and temperamental differences (five domains: hyperthymic, dysthymic, cyclothymic, irritable, and anxious) via the Temperament Evaluation Questionnaire (TEMPS-A)45–46.

Earlier version: Additionally, we incorporated a selection of questionnaires used by Gillan et al. (2016)46, which assess three factors—Anxious-depression, Compulsive behavior and intrusive thought, and Social withdrawal—linked to goal-directed control deficits. These include the Apathy Evaluation Scale47 (AES), Eating Attitudes Test48 (EAT), Alcohol Use Disorders Identification Test49 (AUDIT), Barratt Impulsiveness Scale50 (BIS), Liebowitz Social Anxiety Scale51 (LSAS), Self- Rating Depression Scale52 (SDS), Trait Subscale of the State-Trait Anxiety Inventory for Adults53 (STAI; Spielberger, 1983), and the revised Obsessive-Compulsive Inventory54 (OCI-R). We utilized abbreviated versions of these questionnaires following the methodology of Wise and Dolan (2020)55, who refined the questionnaires using a classifier to retain the items most contributory to variance in the three factors without significant information loss.

Revised version: Additionally, we incorporated a selection of questionnaires used by Gillan et al. (2016)47, which assess three factors—Anxious-depression, Compulsive behavior and intrusive thought, and Social withdrawal—linked to goal-directed control deficits. These include the Apathy Evaluation Scale48 (AES), Eating Attitudes Test49 (EAT), Alcohol Use Disorders Identification Test50 (AUDIT), Barratt Impulsiveness Scale51 (BIS), Liebowitz Social Anxiety Scale52 (LSAS), Self-Rating Depression Scale53 (SDS), Trait Subscale of the State-Trait Anxiety Inventory for Adults54 (STAI; Spielberger, 1983), and the revised Obsessive-Compulsive Inventory55 (OCI-R). We utilized abbreviated versions of these questionnaires following the methodology of Wise and Dolan (2020)56, who refined the questionnaires using a classifier to retain the items most contributory to variance in the three factors without significant information loss.

Earlier version: As preregistered, we aggregated the scores of the selected items from each questionnaire independently for inclusion in our CCA. Questionnaires were scored following the corresponding scoring guide, but in such a way that larger scores always indicate greater psychopathology. Based on best practice recommendations56–57, participants were screened for inattentive responding with infrequency items (‘I competed in the 1917 Summer Olympic Games’) which resulted in the removal of 11 participants (Initial: 6, Replication: 5).

Revised version: As preregistered, we aggregated the scores of the selected items from each questionnaire independently for inclusion in our CCA. Questionnaires were scored following the corresponding scoring guide, but in such a way that larger scores always indicate greater psychopathology. Based on best practice recommendations57–58, participants were screened for inattentive responding with infrequency items (‘I competed in the 1917 Summer Olympic Games’) which resulted in the removal of 11 participants (Initial: 6, Replication: 5).

Directly using the scores from questionnaires of varying lengths in the CCA can give undue influence to short questionnaires, whose scores are likely more noisy. To mitigate this problem, we adjusted the standard deviation of each questionnaire’s score proportionally to the square root of its number of items before conducting the CCA (in accordance with the formula for standard error). After performing the CCA on these weighted scores and normalized model parameters, the resulting canonical loadings were transformed back to their original standardized scale, ensuring that the reported loadings accurately reflected the original, unweighted scale of the questionnaire scores.

Canonical Correlation Analysis (CCA)

Earlier version: Canonical correlation analysis was implemented using the robust CCA method based on projection pursuit (’ccAPP’ package;)58. We searched for the single canonical dimension that maximizes the Spearman correlation between the model parameters and psychopathology scores. Significance testing was conducted using a two-tailed permutation test with 1,000 iterations. In each iteration, the model parameters were shuffled relative to the psychopathology scores, and the canonical correlation was recalculated. The observed canonical correlation was then compared to this null distribution to obtain a p-value. To evaluate the contribution of the elasticity prior parameter (γelasticity) to the parameter-psychopathology correlation, we performed an additional permutation test. In this test, only the elasticity bias estimates (γelasticity) were shuffled, while all other model parameters and self-report scores remained constant.

Revised version: Canonical correlation analysis was implemented using the robust CCA method based on projection pursuit (’ccAPP’ package;)59. We searched for the single canonical dimension that maximizes the Spearman correlation between the model parameters and psychopathology scores. Significance testing was conducted using a two-tailed permutation test with 1,000 iterations. In each iteration, the model parameters were shuffled relative to the psychopathology scores, and the canonical correlation was recalculated. The observed canonical correlation was then compared to this null distribution to obtain a p-value. To evaluate the contribution of the elasticity prior parameter (𝛾elasticity) to the parameter-psychopathology correlation, we performed an additional permutation test. In this test, only the elasticity bias estimates (𝛾elasticity) were shuffled, while all other model parameters and self-report scores remained constant.

Earlier version: To assess how the identified psychopathological profile manifests in raw task measures, we calculated a composite psychopathology score for each participant. This score was derived by multiplying individual participant scores by CCA-derived loadings, representing their position on the psychopathology spectrum. We then used this composite score as a predictor (βcomposite) in a linear regression model to forecast the difference in average opt-in and extra ticket behavior between conditions where controllability was high and inelastic versus high and elastic.

Revised version: To assess how the identified psychopathological profile manifests in raw task measures, we calculated a composite psychopathology score for each participant. This score was derived by multiplying individual participant scores by CCA-derived loadings, representing their position on the psychopathology spectrum. We then used this composite score as a predictor (𝛽𝑐omposite) in a linear regression model to forecast the difference in average opt-in and extra ticket behavior between conditions where controllability was high and inelastic versus high and elastic.